# Homework 1 Assignment


The places where you have enter code are marked with `# YOUR CODE HERE`.

## Question 1 (8 points)

Complete the function `OR3` that implements the 3-qubit to 1-qubit $OR$ operator. This operator acts on 4 qubits in the state $|q_0,q_1,q_2,q_3\rangle$ and maps it to $|q_0,q_1,q_2,q_3 \oplus (q_0 \text{ or } q_1 \text{ or } q_2)  \rangle$.

The function `OR3` has

- Inputs: None
- Returns: A `QuantumCircuit` object with 4 qubits and 1 classical bit that implements the $OR$ operator as described above.

Don't add any measurements to the returned circuit.

```python
!pip install qiskit==0.43.2
```

Output:
```
Collecting qiskit==0.43.2
  Downloading qiskit-0.43.2.tar.gz (9.1 kB)
  Installing build dependencies ... [?25l[?25hdone
  Getting requirements to build wheel ... [?25l[?25hdone
  Installing backend dependencies ... [?25l[?25hdone
  Preparing metadata (pyproject.toml) ... [?25l[?25hdone
Collecting qiskit-terra==0.24.1 (from qiskit==0.43.2)
  Downloading qiskit_terra-0.24.1-cp310-cp310-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (5.9 MB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m5.9/5.9 MB[0m [31m28.1 MB/s[0m eta [36m0:00:00[0m
[?25hCollecting qiskit-aer==0.12.1 (from qiskit==0.43.2)
  Downloading qiskit_aer-0.12.1-cp310-cp310-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (12.8 MB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m12.8/12.8 MB[0m [31m2.9 MB/s[0m eta [36m0:00:00[0m
[?25hCollecting qiskit-ibmq-provider==0.20.2 (from qiskit==0.43.2)
  Downloading qiskit_ibmq_provider-0.20.2-py3-none-any.whl (241 kB)
[2K     [90m
```

<!-- grade_id: cell-517fe01d3e0bd3be, points: 0, grade: false, solution: true -->
```python
from qiskit import QuantumRegister, ClassicalRegister, QuantumCircuit
from qiskit.circuit.library.standard_gates import XGate
from qiskit.circuit.library import MCXGate
from qiskit_aer import AerSimulator


def OR3():
    circuit = QuantumCircuit(4,1)
    # Do not modify anything above this line

    # YOUR CODE HERE
    circuit.cx(0, 3)  # If q0 is 1, set q3
    circuit.cx(1, 3)  # If q1 is 1, set q3
    circuit.cx(2, 3)  # If q2 is 1, set q3

    # Reset q3 if necessary using Toffoli gates
    circuit.ccx(0, 1, 3)  # If q0 and q1 are both 1, reset q3 (double count)
    circuit.ccx(0, 2, 3)  # If q0 and q2 are both 1, reset q3 (double count)
    circuit.ccx(1, 2, 3)  # If q1 and q2 are both 1, reset q3 (double count)

    mcx_gate = MCXGate(3)
    circuit.append(mcx_gate, [0,1,2,3])



    # Do not modify anything below this line
    return circuit
```

In the following cell, there is some code that will help you make sure your solution is correct. You can modify this code as you wish. You may also create more cells and write more code to test your own solution.

```python
def test(input):
    circuit = QuantumCircuit(4,1)

    # implement the input state with a suitable use of X gates.
    for i in range(3):
        if input[i]=='1':
            circuit.x(i)

    # Apply the OR3 gate you created
    circuit.compose(OR3(), range(4), [0], inplace=True)

    circuit.measure(3,0)

    job = AerSimulator().run(circuit,shots=100)
    return job.result().get_counts(circuit)

print(test('000'))
print(test('001'))
print(test('010'))
print(test('011'))
print(test('100'))
print(test('101'))
print(test('110'))
print(test('111'))

```

Output:
```
{'0': 100}
{'1': 100}
{'1': 100}
{'1': 100}
{'1': 100}
{'1': 100}
{'1': 100}
{'1': 100}
```

# How autograding will work in this workshop
In the next cell, there are some tests (the `assert` statements) that will be used to *grade* if your solution is correct. If executing the next cell results in an error, your implementation is incorrect, and you will not get any marks.

There are also some hidden tests, not available to you, that will test your solution, once your submit this file back on Canvas. If any of those fail, you will not receive any marks. Otherwise, you will gain full marks on this question.

<!-- grade_id: cell-12b1e3e97b88002b, points: 2, grade: true, solution: false -->
```python
def Grade_test(input):
    circuit = QuantumCircuit(4,1)

    for i in range(3):
        if input[i]=='1':
            circuit.x(i)

    circuit.compose(OR3(), range(4), [0], inplace=True)
    circuit.measure(3,0)

    job = AerSimulator().run(circuit,shots=100)
    return job.result().get_counts(circuit)

counts = Grade_test('100')
assert counts['1']==100
counts2 = Grade_test('000')
assert counts2['0']==100
```

<!-- grade_id: cell-74f92b92e2a40b2e, points: 6, grade: true, solution: false -->
```python
# hidden tests in this cell will be used for grading.
```